<a href="https://colab.research.google.com/github/GabrielZCipriano/GabrielZCipriano/blob/main/Aula_03_Exerc%C3%ADcios_Python_for_ML_Gabriel_Zanette_Cipriano.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>


# Machine Learning
## Lesson: Python for Machine Learning

This lesson introduces __python__, focusing on __machine learning__.


### Professor:

<img src="https://www.robocup.org.br/wp/wp-content/uploads/2020/02/foto_alexandre-e1581877387559.jpg" width="100" style="float: left; margin-right: 5px;" border="10px" />

  __Prof. Dr. Alexandre da Silva Simões__ <br>
  Department of Control and Automation (DECA) <br>
  Institute of Science and Technology of Sorocaba (ICTS) <br>
  São Paulo State University (Unesp) <br>
  www.sorocaba.unesp.br/professor/assimoes

<br/>


```python
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Mounted at /content/drive
```

## Exercise 1 (5,0 points): Basic math operations and functions

Your __goals__ in this exercise are:

1. (4,0 points) Define four functions (names: mysum, mysub, mymult, mydiv) that will perform the four basic operations over their two arguments
2. (1,0 point ) Define a function (name: myconcat) that will concatenate its four text arguments.

OBS:
* Assume that input numbers can be floats.
* Numerical outputs must be floats with exactly 1 decimal place


<!-- grade_id: my_sum, points: 0, grade: false, solution: true -->
```python

def mysum(a,b):
    return round(a+b,1)

```

<!-- grade_id: my_sum_test, points: 1, grade: true, solution: false -->
```python
assert mysum(0,0) == 0
assert mysum(1,1) == 2
assert mysum(2,12) == 14
assert mysum(1.3,2.4) == 3.7
assert mysum(8.200001,5) == 13.2
```

<!-- grade_id: my_sub, points: 0, grade: false, solution: true -->
```python

def mysub(a,b):
    return round(a-b,1)

```

<!-- grade_id: my_sub_test, points: 1, grade: true, solution: false -->
```python
assert mysub(0,0) == 0
assert mysub(1,1) == 0
assert mysub(10,2) == 8
assert mysub(2,12) == -10
assert mysub(2.3,1.4) == 0.9
assert mysub(1.4999999,2.3) == -0.8
```

<!-- grade_id: my_mult, points: 0, grade: false, solution: true -->
```python

def mymult(a,b):
    return round(a*b,1)

```

<!-- grade_id: my_mult_test, points: 1, grade: true, solution: false -->
```python
assert mymult(0,0) == 0
assert mymult(1,1) == 1
assert mymult(10,2) == 20
assert mymult(2,-12) == -24
assert mymult(2.3,1.4) == 3.2
```

<!-- grade_id: my_div, points: 0, grade: false, solution: true -->
```python

def mydiv(a,b):
    return(round(a/b,1))

```

<!-- grade_id: my_div_test, points: 1, grade: true, solution: false -->
```python
assert mydiv(1,1) == 1
assert mydiv(10,2) == 5
assert mydiv(10,4) == 2.5
assert mydiv(1.4,2.3) == 0.6
```

<!-- grade_id: my_concat, points: 0, grade: false, solution: true -->
```python
def myconcat(a,b,c,d):
    return(a+b+c+d)
```

<!-- grade_id: my_concat_teset, points: 1, grade: true, solution: false -->
```python
assert myconcat('c','a','s','a') == 'casa'
assert myconcat('a','ma','re','lo') == 'amarelo'
assert myconcat('a',' ','h','a') == 'a ha'
```

## Exercise 2 (1,0 point): basic concepts

Your __goal__ in this exercise is to create a function that returns a tuple with both area and perimeter of a circle. The function receives the circle diameter as a parameter. Consider pi=3.14. Solution must be rounded with 2 decimals.


<!-- grade_id: circle, points: 0, grade: false, solution: true -->
```python
def my_circle(d):
    pi = 3.14
    r = d/2
    a = pi*r**2
    p = 2*pi*r
    return(round(a,2), round(p,2))
```

<!-- grade_id: circle-test, points: 1, grade: true, solution: false -->
```python
assert my_circle(1) == (0.79, 3.14)
assert my_circle(10) == (78.5, 31.4)
assert my_circle(4.33) == (14.72, 13.6)
```

## Exercise 3 (1,0 point): Data types

Your __goal__ in this exercise is to write a program which receives a single parameter which is a string with the full name of a person (with spaces) and that returns this name in reverse order.

Example: João da Silva -> Silva da João

```python
palavras = ["Olá", "mundo"]
concatenada = " ".join(palavras)
print(concatenada)
```

Output:
```
Olá mundo
```

<!-- grade_id: dompedro, points: 0, grade: false, solution: true -->
```python
def revert_name(full_name):
    a = full_name.split(' ')
    b = list(reversed(a))
    c = " ".join(b)
    return(c)
    round(avg,2)
```

<!-- grade_id: dompedro-answer, points: 1, grade: true, solution: false -->
```python
assert revert_name('Alexandre da Silva Simões') == 'Simões Silva da Alexandre'
assert revert_name('Albus Percival Wulfric Brian Dumbledore') == 'Dumbledore Brian Wulfric Percival Albus'
assert revert_name('Pedro de Alcântara João Carlos Leopoldo Salvador Bibiano Francisco Xavier de Paula Leocádio Miguel Gabriel Rafael Gonzaga') == 'Gonzaga Rafael Gabriel Miguel Leocádio Paula de Xavier Francisco Bibiano Salvador Leopoldo Carlos João Alcântara de Pedro'
```

## Exercise 4 (1,0 point): list of arguments

Your __goal__ in this exercise is to build a function that receives a list with an arbitrary number of heights of people and returns the average of these heights. The output must be rounded with 2 decimal places.



<!-- grade_id: average, points: 0, grade: false, solution: true -->
```python
def compute_average(my_list):
    soma = sum(my_list)
    tam = len(my_list)
    avg = soma/tam
    return(round(avg,2))
```

<!-- grade_id: average-test, points: 1, grade: true, solution: false -->
```python
assert compute_average([2,4,6]) == 4.00
assert compute_average([2]) == 2
assert compute_average([3, 4, 5, 7, 2322, 32.233, 4566.2093]) == 991.35
```


## Exercise 5 (2,0 points): If-then structure / data conversion

Your __goal__ in this exercise is to assign the final grade (A, B, C or F) to students based on 3 elements:
* number of classes (nc)
* number of presences (np) and
* tests average (av).

The criteria for approval are:

* If the presence of the student in the classes is smaller than 70%, he/she failed (F)
* If the student attended the 70% presence criteria, his/her final grade will be:
    * A if his/her exam grade is equal or bigger than 8,5
    * B if his/her exam grade is smaller than 8,5 and equal or bigger than 7,0
    * C if his/her exam grade is smaller than 7,0 and equal or bigger than 5,0
* The student will fail (F) if his/her exam grade is smaller than 5,0.


<!-- grade_id: if-then, points: 0, grade: false, solution: true -->
```python
from re import A
def grade(nc,np,av):
    if av >= 8.5 and np >= 7:
        nota = "A"
    elif av < 8.5 and av >= 7 and np >= 7:
        nota = "B"
    elif av < 7 and av >= 5 and np >= 7:
        nota = "C"
    else :
        nota ="F"
    return(nota)
```

<!-- grade_id: if-then_test, points: 2, grade: true, solution: false -->
```python
assert grade(10,9,2) == 'F'
assert grade(10,9,4.99) == 'F'
assert grade(10,9,5) == 'C'
assert grade (10,9,7) == 'B'
assert grade (10,9,8.5) == 'A'
assert grade (10,9,10) == 'A'
assert grade (10,7,7) == 'B'
assert grade (10,6.9,10) == 'F'
assert grade (10,2,10) == 'F'
```

## Exercise 6 (5,0 points): Data files read and manipulation

In the file "zoo.dat" (available in the same directory as this notebook) you will find a database about animals and their characteristics. This database contains 17 attributes (rows) of 101 animals (lines). Dataset source: UCI Machine Learning Repository (https://archive.ics.uci.edu/ml/datasets/Zoo)

Attribute Information:

1. animal name: Text
2. hair: Boolean
3. feathers: Boolean
4. eggs: Boolean
5. milk: Boolean
6. airborne: Boolean
7. aquatic: Boolean
8. predator: Boolean
9. toothed: Boolean
10. backbone: Boolean
11. breathes: Boolean
12. venomous: Boolean
13. fins: Boolean
14. legs: Numeric (set of values: {0,2,4,5,6,8})
15. tail: Boolean
16. domestic: Boolean
17. catsize: Boolean
18. type: Numeric (integer values in range [1,7])

OBS:
* Number of Instances: 101
* Number of Attributes: 17
* The file has 2 instances of "frog" and strange animals like "girl" and "vampire"...

Your __goal__ in this exercise is:
1. To load all data stored in the file in an ordered way to a matrix;
2. Return properly the data for two distincts types of question:\
   2.1 When receiving parameters that refers to an animal (ex: zoo('animal', 'elephant')) the function should return the line with all information about this aninal, that is, a vector with 18 entries like: ['elephant','1','0','0','1','0','0','0','1','1','1','0','0','4','1','0','1','1'] \
   2.2 Whe receiving parameters about characteristics of animals (ex: zoo('eggs','1')) the function should return a vector (with an arbitrary number of entries) containing the names of the animals that match the requested conditions (in this case, those who lays eggs):['bass','carp','catfish','chicken','chub','clam','crab','crayfish','crow','dogfish','dove','duck','flamingo','flea','frog','frog','gnat','gull','haddock','hawk','herring','honeybee','housefly','kiwi','ladybird','lark','lobster','moth','newt','octopus','ostrich','parakeet','penguin','pheasant','pike','piranha','pitviper','platypus','rhea','seahorse','seawasp','skimmer','skua','slowworm','slug','sole','sparrow','starfish','stingray','swan','termite','toad','tortoise','tuatara','tuna','vulture','wasp','worm','wren']




<!-- grade_id: zoo, points: 0, grade: false, solution: true -->
```python
def zoo(field, value):
    columns = {
        'animal': 0,
        'hair': 1,
        'feathers': 2,
        'eggs': 3,
        'milk': 4,
        'airborne': 5,
        'aquatic': 6,
        'predator': 7,
        'toothed': 8,
        'backbone': 9,
        'breathes': 10,
        'venomous': 11,
        'fins': 12,
        'legs': 13,
        'tail': 14,
        'domestic': 15,
        'catsize': 16,
        'type': 17
    }
    with open("/content/drive/MyDrive/AMRDC/Exercícios/zoo.dat", "r") as f:
      lines = f.readlines()
    data = []

    if field == 'animal':
      for line in lines:
        animal = line.strip().split(',')[0]

        if animal == value:
          return(line.strip().split(','))

    else:
      for line in lines:
        characteristics = line.strip().split(',')[columns[field]]

        if characteristics == value and (line.strip().split(',')[0] != 'aardvark'):
          data.append(line.strip().split(',')[0])
      return data
    f.close()

print(zoo('animal', 'octopus'))
print(zoo('tail', '0'))
print(zoo('legs', '2'))
print(zoo('milk', '1'))
```

Output:
```
['octopus', '0', '0', '1', '0', '0', '1', '1', '0', '0', '0', '0', '0', '8', '0', '0', '1', '7']
['bear', 'cavy', 'clam', 'crab', 'crayfish', 'flea', 'frog', 'frog', 'girl', 'gnat', 'gorilla', 'honeybee', 'housefly', 'ladybird', 'lobster', 'moth', 'octopus', 'seal', 'seawasp', 'slug', 'starfish', 'termite', 'toad', 'wasp', 'worm']
['chicken', 'crow', 'dove', 'duck', 'flamingo', 'fruitbat', 'girl', 'gorilla', 'gull', 'hawk', 'kiwi', 'lark', 'ostrich', 'parakeet', 'penguin', 'pheasant', 'rhea', 'sealion', 'skimmer', 'skua', 'sparrow', 'squirrel', 'swan', 'vampire', 'vulture', 'wallaby', 'wren']
['antelope', 'bear', 'boar', 'buffalo', 'calf', 'cavy', 'cheetah', 'deer', 'dolphin', 'elephant', 'fruitbat', 'giraffe', 'girl', 'goat', 'gorilla', 'hamster', 'hare', 'leopard', 'lion', 'lynx', 'mink', 'mole', 'mongoose', 'opossum', 'oryx', 'platypus', 'polecat', 'pony', 'porpoise', 'puma', 'pussycat', 'raccoon', 'reindeer', 'seal', 'sealion', 'squirrel', 'vampire', 'vole', 'wallaby', 'wolf']
```

<!-- grade_id: zoo_test, points: 5, grade: true, solution: false -->
```python
#print all 17 characteristics of the octopus:
assert zoo('animal', 'octopus') == ['octopus', '0', '0', '1', '0', '0', '1', '1', '0', '0', '0', '0', '0', '8', '0', '0', '1', '7']
#print the list of all animals that don't have tail:
assert zoo('tail', '0') == ['bear','cavy','clam','crab','crayfish','flea','frog','frog','girl','gnat','gorilla','honeybee','housefly','ladybird','lobster','moth','octopus','seal','seawasp','slug','starfish','termite','toad','wasp','worm']
#print the list of all bipedal animals
assert zoo('legs', '2') == ['chicken','crow','dove','duck','flamingo','fruitbat','girl','gorilla','gull','hawk','kiwi','lark','ostrich','parakeet','penguin','pheasant','rhea','sealion','skimmer','skua','sparrow','squirrel','swan','vampire','vulture','wallaby','wren']
#print the list of animals that are mammals
assert zoo('milk','1') == ['antelope','bear','boar','buffalo','calf','cavy','cheetah','deer','dolphin','elephant','fruitbat','giraffe','girl','goat','gorilla','hamster','hare','leopard','lion','lynx','mink','mole','mongoose','opossum','oryx','platypus','polecat','pony','porpoise','puma','pussycat','raccoon','reindeer','seal','sealion','squirrel','vampire','vole','wallaby','wolf']
```